---
id: 66f2836c459cfb16ae76f24f
title: Construye un generador de pirámide
challengeType: 25
dashedName: lab-pyramid-generator
---

# --description--

Complete los relatos de usuario a continuación y obtenga la aprobación de todas las pruebas para completar el laboratorio.

**Instrucciones:**

1. Debe tener una función llamada `pyramid` que reciba tres argumentos.
1. El primer argumento debe ser una cadena que represente el carácter del patrón que se repetirá en su pirámide.
1. El segundo argumento debe ser un número entero que represente el número de filas de la pirámide.
1. El tercer argumento debe ser un valor booleano.
1. La función `pyramid` debe devolver una cadena en la que el carácter del patrón se repite y se organiza para formar una pirámide con el vértice hacia arriba cuando el tercer argumento es `false`.
1. Cuando el tercer argumento es `true`, la pirámide debe tener el vértice hacia abajo.
1. La fila del vértice debe tener un único carácter de patrón y en las siguientes filas debe agregar dos caracteres más del patrón que la fila anterior.
1. Cada fila debe comenzar con una cantidad de espacios suficiente para colocar el carácter central de cada fila en la misma columna.
1. La pirámide debe comenzar y terminar con un carácter de nueva línea.

Por ejemplo, el llamado a  `pyramid("o", 4, false)` debería dar este resultado:

```js

   o
  ooo
 ooooo
ooooooo

```

# --hints--

Debes tener una función llamada `pyramid`.

```js
assert.isFunction(pyramid);
```

Su función `pyramid` debe tener tres parámetros.

```js
assert.lengthOf(pyramid, 3);
```

`pyramid("o", 4, false)` debe devolver `"\n   o\n  ooo\n ooooo\nooooooo\n"`.

```js
assert.equal(pyramid("o", 4, false), "\n   o\n  ooo\n ooooo\nooooooo\n")
```

`pyramid("p", 5, true)` debe devolver `"\nppppppppp\p ppppppp\n  ppppp\n   ppp\n    p\n"`.

```js
assert.equal(pyramid("p", 5, true), "\nppppppppp\n ppppppp\n  ppppp\n   ppp\n    p\n")
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function pyramid(char, count, isInverted) {
    const rows = []
    for (let i = 1; i <= count; i++) {
        if (isInverted) {
            rows.unshift(" ".repeat(count - i) + char.repeat(2 * i - 1))
        } else {
            rows.push(" ".repeat(count - i) + char.repeat(2 * i - 1))
        }        
    }
    return "\n" + rows.join("\n") + "\n";
}

console.log(pyramid("#", 10, false))

```
